```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.2.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 6.2: Assessing a Models</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

**Reading**: 
* [Assessing a Model](https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html)

Please complete this notebook by filling in the cells provided.

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## Assessing Jade's Models

### Games with Jade

Our friend Jade comes over and asks us to play a game with her. The game works like this: 

> We will draw randomly with replacement from a simplified 13 card deck with 4 face cards (A, J, Q, K), and 9 numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10). If we draw cards with replacement 13 times, and if the number of face cards is greater than or equal to 4, we lose.
> 
> Otherwise, Jade wins.

We play the game once and we lose, observing 8 total face cards. We are angry and accuse Jade of cheating! Jade is adamant, however, that the deck is fair.

Jade's model claims that there is an equal chance of getting any of the cards (A, 2, 3, 4, 5, 6, 7, 8, 9, 10, J, Q, K), but we do not believe her. We believe that the deck is clearly rigged, with face cards (A, J, Q, K) being more likely than the numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10).

#### Task 01 📍

Assign `deck_model_probabilities` to a two-item array containing the chance of drawing a face card as the first element, and the chance of drawing a numbered card as the second element under Jade's model. Since we're working with probabilities, make sure your values are between 0 and 1. 


_Points:_ 3

```python
deck_model_probabilities = ...
deck_model_probabilities
```

```python
grader.check("task_01")
```

#### Task 02 📍

We believe Jade's model is incorrect. In particular, we believe there to be a  larger chance of getting a face card. Which of the following statistics can we use during our simulation to test between the model and our alternative? Assign `statistic_choice` to the correct answer. 

1. The actual number of face cards we get in 13 draws
2. The distance (absolute value) between the actual number of face cards in 13 draws and the expected number of face cards in 13 draws (4)
3. The expected number of face cards in 13 draws (4)



_Points:_ 2

```python
statistic_choice = ...
statistic_choice
```

```python
grader.check("task_02")
```

#### Task 03 📍

Define the function `deck_simulation_and_statistic`, which, given a sample size and an array of model proportions (like the one you created in Task 1), returns the number of face cards in one simulation of drawing a card under the model specified in `model_proportions`. 

As you form your response, think about how you can use the function `sample_proportions`. 


_Points:_ 1

```python
def deck_simulation_and_statistic(sample_size, model_proportions):
    ...

deck_simulation_and_statistic(13, deck_model_probabilities)
```

```python
grader.check("task_03")
```

#### Task 04 📍

Use your function from above to simulate the drawing of 13 cards 5000 times under the proportions that you specified in Task 1. Keep track of all of your statistics in `deck_statistics`. 


_Points:_ 2

```python
repetitions = 5000 
...

deck_statistics
```

```python
grader.check("task_04")
```

Let’s take a look at the distribution of simulated statistics.

```python
#Draw a distribution of statistics 
Table().with_column('Deck Statistics', deck_statistics).hist()
```

#### Task 05 📍🔎

<!-- BEGIN QUESTION -->

Given your observed value, do you believe that Jade's model is reasonable, or is our alternative more likely? Explain your answer using the distribution drawn in the previous problem. 


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 3)
```
>>> 
>>> len(deck_model_probabilities) == 2
True
```
failure_message: ❌ deck_model_probabilities should have a length 2
success_message: ✅ deck_model_probabilities has length 2
```
>>> 
>>> sum(deck_model_probabilities) == 1
True
```
failure_message: ❌ The sum of the values in deck_model_probabilities should be 1.
success_message: ✅ The sum of the values in deck_model_probabilities is 1.


Tests for task_02 (points: 2)
```
>>> 
>>> statistic_choice in [1, 2, 3]
True
```
failure_message: ❌ statistic_choice should name the value 1, 2, or 3.
success_message: ✅ statistic_choice represents a possible value.


Tests for task_03 (points: 1)
```
>>> 
>>> # The statistic should be between 0 and 13 face cards for a sample size of 13
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> 0 <= num_face <= 13
True
```
failure_message: ❌ The statistic should be between 0 and 13 face cards for a sample size of 13.
success_message: ✅ The calculated value of the statistic is a possible value.


Tests for task_04 (points: 2)
```
>>> 
>>> len(deck_statistics) == repetitions
True
```
failure_message: ❌ There should be exactly as many elements in deck_statistics as the number 'repetitions'.
success_message: ✅ There is exactly as many elements in deck_statistics as the number 'repetitions'.
```
>>> 
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```
failure_message: ❌ Each element of deck_statistics should be between 0 and 13 inclusive.
success_message: ✅ Each element of deck_statistics is between 0 and 13 inclusive
