```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Lab 10: Regression

## References

* [Sections 15.0 - 15.6 of the Textbook](https://inferentialthinking.com/chapters/15/Prediction.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Python Quick Reference](https://ccsf-math-108.github.io/materials-fa23/resources/quick_reference.html)

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- For all tasks indicated with a 🔎 that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) Please don't just share answers, though.
- View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.

Run the following cell to set up the lab, and make sure you run the cell at the top of the notebook that initializes Otter.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## How Faithful is Old Faithful? 

Old Faithful is a geyser in Yellowstone National Park that is famous for eruption on a fairly regular schedule. Run the cell below to see Old Faithful in action!

```python
from IPython.display import YouTubeVideo
YouTubeVideo("wE8NDuzt8eg")
```

Some of Old Faithful's eruptions last longer than others.  Whenever there is a long eruption, it is usually followed by an even longer wait before the next eruption. If you visit Yellowstone, you might want to predict when the next eruption will happen, so that you can see the rest of the park instead of waiting by the geyser.
 
Today, we will use a dataset on eruption durations and waiting times to see if we can make such predictions accurately with linear regression.

The dataset has one row for each observed eruption.  It includes the following columns:
- `duration`: Eruption duration, in minutes
- `wait`: Time between this eruption and the next, also in minutes

Run the next cell to load the dataset.

```python
faithful = Table.read_table("faithful.csv")
faithful
```

We would like to use linear regression to make predictions, but that won't work well if the data aren't roughly linearly related.  To check that, we should look at the data.

### Task 01 📍🔎

Make a scatter plot of the data.  It's conventional to put the column we want to predict on the vertical axis and the other column on the horizontal axis.


```python
...
```

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

Are eruption duration and waiting time roughly linearly related based on the scatter plot above? Is this relationship positive?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

We're going to continue with the assumption that they are linearly related, so it's reasonable to use linear regression to analyze this data.

We'd next like to plot the data in standard units. If you don't remember the definition of standard units, textbook section [14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html#standard-units) might help!

### Task 03 📍

Compute the mean and standard deviation of the eruption durations and waiting times.  **Then** create a table called `faithful_standard` containing the eruption durations and waiting times in standard units.  The columns should be named `duration (standard units)` and `wait (standard units)`.


```python

duration_mean = ...
duration_std = ...
wait_mean = ...
wait_std = ...

faithful_standard = Table().with_columns(
    "duration (standard units)", ...,
    "wait (standard units)", ...)
faithful_standard

```

```python
grader.check("task_03")
```

### Task 04 📍🔎

Plot the data again, but this time in standard units.


```python
...
```

You'll notice that this plot looks the same as the last one!  However, the data and axes are scaled differently.  So it's important to read the ticks on the axes.

### Task 05 📍

Among the following numbers, which would you guess is closest to the correlation between eruption duration and waiting time in this dataset?

1. -1
2. 0
3. 1

Assign `correlation` to the number corresponding to your guess.


```python
correlation = ...
```

```python
grader.check("task_05")
```

### Task 06 📍

Compute the correlation `r`.  

*Hint:* Use `faithful_standard`.  Section [15.1](https://www.inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r) explains how to do this.



```python
r = ...
r
```

```python
grader.check("task_06")
```

## The regression line

Recall that the correlation is the **slope of the regression line when the data are put in standard units**.

The next cell plots the regression line in standard units:

$$\text{waiting time in standard units} = r \times \text{eruption duration in standard units}.$$

Then, it plots the data in standard units again, for comparison.

```python
def plot_data_and_line(dataset, x, y, point_0, point_1):
    """Makes a scatter plot of the dataset, along with a line passing through two points."""
    dataset.scatter(x, y, label="data")
    xs, ys = zip(point_0, point_1)
    plt.plot(xs, ys, label="regression line")
    plt.legend(bbox_to_anchor=(1.5,.8))

plot_data_and_line(faithful_standard, 
                   "duration (standard units)", 
                   "wait (standard units)", 
                   [-2, -2*r], 
                   [2, 2*r])
```

How would you take a point in standard units and convert it back to original units?  We'd have to "stretch" its horizontal position by `duration_std` and its vertical position by `wait_std`. That means the same thing would happen to the slope of the line.

Stretching a line horizontally makes it less steep, so we divide the slope by the stretching factor.  Stretching a line vertically makes it more steep, so we multiply the slope by the stretching factor.

### Task 07 📍

Calculate the slope of the regression line in original units, and assign it to `slope`.

(If the "stretching" explanation is unintuitive, consult section [15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line) in the textbook.)


```python
slope = ...
slope
```

```python
grader.check("task_07")
```

We know that the regression line passes through the point `(duration_mean, wait_mean)`.  You might recall from high-school algebra that the equation for the line is therefore:

$$\text{waiting time} - \verb|wait_mean| = \texttt{slope} \times (\text{eruption duration} - \verb|duration_mean|)$$

The rearranged equation becomes:

$$\text{waiting time} = \texttt{slope} \times \text{eruption duration} + (- \texttt{slope} \times \verb|duration_mean| + \verb|wait_mean|)$$

### Task 08 📍

Calculate the intercept in original units and assign it to `intercept`.


```python
intercept = ...
intercept
```

```python
grader.check("task_08")
```

## Investigating the regression line

The slope and intercept tell you exactly what the regression line looks like.  To predict the waiting time for an eruption, multiply the eruption's duration by `slope` and then add `intercept`.

### Task 09 📍🔎 

Compute the predicted waiting time for an eruption that lasts 2 minutes, and for an eruption that lasts 5 minutes.


```python
two_minute_predicted_waiting_time = ...
five_minute_predicted_waiting_time = ...

# Here is a helper function to print out your predictions.
# Don't modify the code below.
def print_prediction(duration, predicted_waiting_time):
    print("After an eruption lasting", duration,
          "minutes, we predict you'll wait", predicted_waiting_time,
          "minutes until the next eruption.")

print_prediction(2, two_minute_predicted_waiting_time)
print_prediction(5, five_minute_predicted_waiting_time)
```

The next cell plots the line that goes between those two points, which is (a segment of) the regression line.

```python
plot_data_and_line(faithful, "duration", "wait", 
                   [2, two_minute_predicted_waiting_time], 
                   [5, five_minute_predicted_waiting_time])
```

### Task 10 📍

Make predictions for the waiting time after each eruption in the `faithful` table.  (Of course, we know exactly what the waiting times were!  We are doing this so we can see how accurate our predictions are.)  Put these numbers into a column in a new table called `faithful_predictions`.  Its first row should look like this:

|duration|wait|predicted wait|
|-|-|-|
|3.6|79|72.1011|

*Hint:* Your answer can be just one line.  There is no need for a `for` loop; use array arithmetic instead.


```python
faithful_predictions = ...
faithful_predictions
```

```python
grader.check("task_10")
```

### Task 11 📍

How close were we?  Compute the *residual* for each eruption in the dataset.  The residual is the actual waiting time minus the predicted waiting time.  Add the residuals to `faithful_predictions` as a new column called `residual` and name the resulting table `faithful_residuals`.

*Hint:* Again, your code will be much simpler if you don't use a `for` loop.


```python
faithful_residuals = ...
faithful_residuals
```

```python
grader.check("task_11")
```

Here is a plot of the residuals you computed.  Each point corresponds to one eruption.  It shows how much our prediction over- or under-estimated the waiting time.

```python
faithful_residuals.scatter("duration", "residual", color="r")
```

There isn't really a pattern in the residuals, which confirms that it was reasonable to try linear regression.  It's true that there are two separate clouds; the eruption durations seemed to fall into two distinct clusters.  But that's just a pattern in the eruption durations, not a pattern in the relationship between eruption durations and waiting times.

## How accurate are different predictions?

Earlier, you should have found that the correlation is fairly close to 1, so the line fits fairly well on the training data.  That means the residuals are overall small (close to 0) in comparison to the waiting times.

We can see that visually by plotting the waiting times and residuals together:

```python
faithful_residuals.scatter("duration", "wait", label="actual waiting time", color="blue")
plt.scatter(faithful_residuals.column("duration"), faithful_residuals.column("residual"), label="residual", color="r")
plt.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time], label="regression line")
plt.legend(bbox_to_anchor=(1.7,.8));
```

However, unless you have a strong reason to believe that the linear regression model is true, you should be wary of applying your prediction model to data that are very different from the training data.

### Task 12 📍

In `faithful`, no eruption lasted exactly 0, 2.5, or 60 minutes.  Using this line, what is the predicted waiting time for an eruption that lasts 0 minutes?  2.5 minutes?  An hour?


```python
zero_minute_predicted_waiting_time = ...
two_point_five_minute_predicted_waiting_time = ...
hour_predicted_waiting_time = ...

print_prediction(0, zero_minute_predicted_waiting_time)
print_prediction(2.5, two_point_five_minute_predicted_waiting_time)
print_prediction(60, hour_predicted_waiting_time)
```

```python
grader.check("task_12")
```

### Task 13 📍🔎

<!-- BEGIN QUESTION -->

For each prediction, state whether you think it's reliable and explain your reasoning. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Divide and Conquer

It appears from the scatter diagram that there are two clusters of points: one for durations around 2 and another for durations between 3.5 and 5. A vertical line at 3 divides the two clusters.

```python
faithful.scatter("duration", "wait", label="actual waiting time", color="blue")
plt.plot([3, 3], [40, 100]);
```

The `standardize` function from lecture appears below, which takes in a table with numerical columns and returns the same table with each column converted into standard units.

```python
def standard_units(any_numbers):
    "Convert any array of numbers to standard units."
    return (any_numbers - np.mean(any_numbers)) / np.std(any_numbers)  

def standardize(t):
    """Return a table in which all columns of t are converted to standard units."""
    t_su = Table()
    for label in t.labels:
        t_su = t_su.with_column(label + ' (su)', standard_units(t.column(label)))
    return t_su
```

### Task 14 📍

Separately compute the regression coefficients *r* for all the points with a duration below 3 **and then** for all the points with a duration above 3. To do so, create a function that computes `r` from a table and pass it two different tables of points, `below_3` and `above_3`.


```python
def reg_coeff(t):
    """Return the regression coefficient for columns 0 & 1."""
    t_su = standardize(t)
    ...

below_3 = ...
above_3 = ...
below_3_r = reg_coeff(below_3)
above_3_r = reg_coeff(above_3)
print("For points below 3, r is", below_3_r, "; for points above 3, r is", above_3_r)
```

```python
grader.check("task_14")
```

### Task 15 📍

Complete the functions `slope_of` and `intercept_of` below. 

When you're done, the functions `wait_below_3` and `wait_above_3` should each use a different regression line to predict a wait time for a duration. The first function should use the regression line for all points with duration below 3. The second function should use the regression line for all points with duration above 3.


```python
def slope_of(t, r):
    """Return the slope of the regression line for t in original units.
    
    Assume that column 0 contains x values and column 1 contains y values.
    r is the regression coefficient for x and y.
    """
    ...

def intercept_of(t, r):
    """Return the slope of the regression line for t in original units."""
    s = slope_of(t, r)
    ...
    
below_3_a = slope_of(below_3, below_3_r)
below_3_b = intercept_of(below_3, below_3_r)
above_3_a = slope_of(above_3, above_3_r)
above_3_b = intercept_of(above_3, above_3_r)

def wait_below_3(duration):
    return below_3_a * duration + below_3_b

def wait_above_3(duration):
    return above_3_a * duration + above_3_b
```

```python
grader.check("task_15")
```

The plot below shows two different regression lines, one for each cluster!

```python
faithful.scatter(0, 1)
plt.plot([1, 3], [wait_below_3(1), wait_below_3(3)])
plt.plot([3, 6], [wait_above_3(3), wait_above_3(6)]);
```

### Task 16 📍

Write a function `predict_wait` that takes a `duration` and returns the predicted wait time using the appropriate regression line, depending on whether the duration is below 3 or greater than (or equal to) 3.


```python
def predict_wait(duration):
    """Return the wait predicted by the appropriate one of the two regression lines above."""
    ...
```

```python
grader.check("task_16")
```

The predicted wait times for each point appear below.

```python
faithful.with_column('predicted', faithful.apply(predict_wait, 'duration')).scatter(0)
```

### Task 17 📍🔎

<!-- BEGIN QUESTION -->

Do you think the predictions produced by `predict_wait` would be more or less accurate than the predictions from the regression line you created in section 2? How could you tell?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submit your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the requirements for a Complete score for this lab assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `grader.check_all()`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
4. Select the menu items `File`, `Save and Export Notebook As...`, and `Html_embed` in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_03 (points: None)
```
>>> abs(sum(faithful_standard.column(0))) <= 1e-08
True
```
failure_message: ❌ Your standardized durations don't seem correct. Try subracting all the durations form the average duration and dividing the result by the standard deviation of the durations.
success_message: ✅ Your standardized durations seem correct.
```
>>> int(round(duration_std)) == 1
True
```
failure_message: ❌ Your duration_std value doesn't seem correct. Try using np.std on the duration column data.
success_message: ✅ Your duration_std value seems correct.
```
>>> int(round(wait_std)) == 14
True
```
failure_message: ❌ Your wait_std value doesn't seem correct. Try using np.std on the wait column data.
success_message: ✅ Your wait_std value seems correct.


Tests for task_05 (points: None)
```
>>> correlation == 3
True
```
failure_message: ❌
success_message: ✅


Tests for task_06 (points: None)
```
>>> '%.2f' % round(r, 2)
'0.90'
```
failure_message: ❌ Your r value doesn't seems correct. Try finding the mean of the product of the standarized durations and wait values.
success_message: ✅ Your r value seems correct.


Tests for task_07 (points: None)
```
>>> (slope * 13 - 100) / 98 <= 0.5
True
```
failure_message: ❌ Your slope value doesn't seem to be correct. Try multiplying the correlation coefficient and the standard deviation of the wait values divided by the standard deviation of the duration values.
success_message: ✅ Your slope value could be correct.
```
>>> np.round(slope, 4) == 10.7296
True
```
failure_message: ❌ Your slope value doesn't seem to be correct. Make sure you are plugging in the SD_x and SD_y in the correct spots!
success_message: ✅ Your slope value seems to be correct.


Tests for task_08 (points: None)
```
>>> 33 < intercept < 34
True
```
failure_message: ❌ Your intercept is not within a reasonable interval. Try subtracting the product of the slope and duration mean from the wait mean.
success_message: ✅ Your intercept is within a reasonable interval.
```
>>> np.round(intercept, 3) == 33.474
True
```
failure_message: ❌ Your intercept value doesn't correct. Try subtracting the product of the slope and duration mean from the wait mean.
success_message: ✅ Your intercept value seems correct.


Tests for task_10 (points: None)
```
>>> set(faithful_predictions.labels) == set(['duration', 'wait', 'predicted wait'])
True
```
failure_message: ❌ The labels in faithful_predictions should be 'duration', 'wait', and 'predicted wait'.
success_message: ✅ The labels in faithful_predictions are correct.
```
>>> abs(1 - np.mean(faithful_predictions.column(2)) / 100) <= 0.35
True
```
failure_message: ❌ Your predicted values in faithful_predictions don't seem correct. Try multiplying the slope times the duration values and adding the intercept.
success_message: ✅ Your predicted values in faithful_predictions seem correct.


Tests for task_11 (points: None)
```
>>> abs(sum(faithful_residuals.column(3))) <= 1e-08
True
```
failure_message: ❌ The residuals in faithful_residuals don't seem correct. Make sure you are subtracting the predicted wait values from the actual wait values.
success_message: ✅ The residuals in faithful_residuals seem correct.


Tests for task_12 (points: None)
```
>>> 12 - zero_minute_predicted_waiting_time * 1.4 / 4 <= 0.35
True
```
failure_message: ❌ The zero_minute_predicted_waiting_time value doesn't seem correct. This should be the intercept.
success_message: ✅ The zero_minute_predicted_waiting_time value seems correct.
```
>>> 2 - two_point_five_minute_predicted_waiting_time / 35 <= 0.4
True
```
failure_message: ❌ The two_point_five_minute_predicted_waiting_time value doesn't seem correct. Multiply the slope times 2.5 and add the intercept.
success_message: ✅ The two_point_five_minute_predicted_waiting_time value seems correct.
```
>>> (26 - hour_predicted_waiting_time / 30) / 10 <= 0.43
True
```
failure_message: ❌ The hour_predicted_waiting_time value doesn't seem correct. Make sure you are using 60 mins!
success_message: ✅ The hour_predicted_waiting_time value seems correct.


Tests for task_14 (points: None)
```
>>> [below_3.num_rows, above_3.num_rows]
[97, 175]
```
failure_message: ❌ The correct number of values are not in below_3 and above_3. Use the where method on the faithful table to find where the duration values are below/above 3.
success_message: ✅ The correct number of values are in below_3 and above_3.
```
>>> np.allclose([below_3_r, above_3_r], [0.290189526493, 0.372782225571])
True
```
failure_message: ❌ The values for below_3_r and above_3_r don't seem correct. Make sure the reg_coeff function returns the mean of the product of the standardized values in column 0 and 1.
success_message: ✅ The values for below_3_r and above_3_r seem correct.


Tests for task_15 (points: None)
```
>>> np.allclose([wait_below_3(1), wait_below_3(3), wait_above_3(3), wait_above_3(6)], [47.90215160574252, 60.60319718202381, 72.96541399053837, 89.2818591974495])
True
```
failure_message: ❌ wait_below_3 and wait_above_3 don't seem to be working correctly. Double check that your slope and intercept functions are using the right formulas for slope and intercept.
success_message: ✅ wait_below_3 and wait_above_3 seem to be working correctly.


Tests for task_16 (points: None)
```
>>> np.allclose([predict_wait(x) for x in [1.5, 2.5, 3.5, 4.5]], [51.077412999812836, 57.42793578795349, 75.68482152502355, 81.12363659399392])
True
```
failure_message: ❌ predict_wait doesn't seem to be working correctly. Try using a if statement to have the function respond differently depending on the duration value.
success_message: ✅ predict_wait seems to be working correctly.
