<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Lab 01 - Jupyter Notebooks</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

## Jupyter Notebooks

<img style="display: block; margin-left: auto; margin-right: auto" src="jupyter-main-logo.svg" alt="Jupyter's main logo" width="100rem;">

This web application is called a Jupyter notebook and is part of [Project Jupyter](https://jupyter.org/). A notebook is a place to host share explanatory content intended for humans usually in association with computational programs and their results. This notebook format supports reproducible computational research where anyone with access to the notebook and the relevant data will more than likely be able to reproduce the work done by the original researchers. Reproducibility is part of the backbone of modern scientific research. Notebook files on a computer end in `.ipynb`. This file extension is a reference to the IPython programming environment developed by [Fernando Pérez](https://en.wikipedia.org/wiki/Fernando_P%C3%A9rez_(software_developer)). You'll notice that these notebooks can be edited and run on several different platforms such as the [CCSF JupyterHub](https://ccsf.cloudbank.2i2c.cloud/), [Google Colab](colab.research.google.com), [VS Code](https://code.visualstudio.com/), etc. The different platforms allow you to interact with the notebooks in slightly different ways, but the main components of a notebook typically persist through those differences.

<img style="display: block; margin-left: auto; margin-right: auto" src="otter-logo-smaller.png" alt="Otter-Grader's logo" width="100rem;">

We will utilize notebooks for lecture demonstrations and assignments such as this one. Typically, you will need to provide explanatory content or code for the associated **Tasks** throughout the notebook. To assist in grading your work, we will use an auto-grading tool called [Otter-Grader](https://otter-grader.readthedocs.io/en/latest). You'll learn about that grader later in this assignment. Let's get started by exploring the various types of cells you will work with in this course.

## Cells

From our point of view, notebooks consist of building blocks called cells. If you navigate through the document, you'll notice that the various cells are outlined by a rectangle. In terms of the basics of cells, you will work with Markdown cells and code cells.

### Markdown Cells

Markdown cells are typically where you will find the content intended for the explanatory material. There could be text, images, etc. in these cells to explain what is going on in the notebook and in the code cells. The text in these cells are formatted with a language called Markdown. You don't need to learn all the details of Markdown, but you might find it helpful to go through the tutorial on [markdowntutorial.com](https://www.markdowntutorial.com/) or reference the [Markdown Cheat Sheet](https://www.markdownguide.org/cheat-sheet/). 

* If you double-click a Markdown cell, then you'll be able to see how it was formatted. Sometimes we lock the Markdown cells in these notebooks, so you won't accidentally modify our explanatory content or instructions. You should have no program editing cells that you create or cells that are associated with the assignment tasks.
* To return back to the formatted version of the cell, select the ▶ button on the toolbar above.

#### Task 1 📍🔎

<!-- BEGIN QUESTION -->

For your first task, edit the following Markdown cell and replace the text `_Type your answer here, replacing this text._`, with your name and a song, movie, book, etc. that you have been enjoying recently. To practice using Markdown, try to format your response as an unordered list with your name in **bold** and the name of the song, movie, book, etc. in _italics_.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Code Cells

Code cells host text associated with some programming language. We will utilize the Python 3 programming language in this course, and you will learn some of the basics of that language throughout the semester. The following cell is a Code cell. You'll notice that Code cells have a different structure than Markdown cells. You can run the Code cell by selecting it and pressing the ▶ button. (As you progress, you'll want to familiarize yourself with a few shortcuts rather than selecting the buttons. For example, to run a Code cell, you can hold down `shift` and press `return`.)

```python
print("Hello, World!")
```

If you ran the above cell, you should notice that next to the Code cell is a set of square brackets with a number inside of it. 
* If a Code cell has not been run, then it will not contain a number `[ ]:`.
* Depending on how long it takes for the code to run, you'll see an asterisk within the brackets while the code is running `[*]:`.
* If it is the first cell to be executed in the notebook, then it will have a 1 in it `[1]:`. 
* The next cell (including the same cell) that is run will have a 2 within the brackets `[2]:`.
* The numbers in the code cells provide us with the sequence in which the Code cells were run. It is important to know which cells have been executed because Code cells ran later in time might depend on Code cells previously executed.

#### Output

Notice that once the above code cell was run, `Hello, World!` appeared below the Code cell as output. The output of a code cell contain images, data tables, etc. produced by the output from running the Python commands.

##### Task 2 📍

In most of the assignments, it will be easy for us to run some checks on your code to see if you are doing things correctly. That auto-grading is done with the help of the Otter-Grader library. Otter runs tests on your code that we create to see how you are doing. Sometimes these tests show you what we are testing for and what the results are, but there will be times when we run hidden tests after you submit your work to help grade for correctness. In most cases, the lab assignments do not have hidden tests, but the homework and project assignments do have hidden tests. For every assignment, you'll run a cell that initializes Otter. Usually, that Code cell will appear at the top of the notebook, but for this first assignment it is below.

For this task:

1. Run the following Code cell that starts with `# Initialize Otter`. Before you run the cell, Python doesn't understand what the work `grader` means, but after you run the cell, Python has an association of the name `grader` with the Otter grader and its configurations for this assignment.
2. Run the grader.check("task_2") cell. 

You should see that there was one test performed and you passed. If you are having troubles, ask for help before moving on. For example, if you find the text `NameError: name 'grader' is not defined` as output and has been shaded red, then that means that you have not run the cell that initializes Otter. 

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

```python
grader.check("task_2")
```

### Adding Cells to Jupyter Notebooks

So far, you have only edited the content of this notebook. There will be times when you want to add your own cells to notebooks. To add a cell, click the + button in the toolbar.  It'll start out as a Code cell.  In order to change the cell type:

1. Click to the cell that you want to change the type of.
2. Click the drop-down box in the toolbar and choose `Markdown`.

### Making Mistakes

Before you begin the next task, we want to remind you that you will make mistakes as you work with these notebooks. Consider the following three pieces of advice:

1. If you find that you have just made a mistake, be careful of compounding the mistakes by trying to fix them. Instead, select `Edit` in the menu bar and select a relevant `Undo` option. 
2.You will create checkpoints in your notebook's history by selecting `File` in the menu bar, then `Save`. You can revert to one of these checkpoints by selecting `File` in the menu bar, followed by `Revert Notebook to Checkpoint`.
2. If you feel that you just want to start your assignment over, then select `File` in the menu bar, followed by `Rename ...`. This keeps your work in your JupyterHub. After that, go back to Canvas and open the link to the notebook from the assignment page. This will create a new copy of the original notebook that we shared with you.

### Task 3 📍🔎

<!-- BEGIN QUESTION -->

For this task,

1. Add a code cell to the notebook and **position it below** this Markdown cell.
2. Type `2 + 2` in the code cell and run it so that `4` appears in its associated output.
3. Add a Markdown cell to the notebook and **position it above** the code cell you created.
4. Type `I'm getting the hang of this! Here is some code:` in the Markdown cell that you added.

You might need to ask for help when it comes to moving the cells around. Don't give up!

<!-- END QUESTION -->

### Edit and Command Modes

Notebooks have two modes. In edit mode, you can edit text in cells. Command mode is used for navigating the notebook using keyboard shortcuts.

- If the notebook is in command mode, press `enter` on the keyboard to change to edit mode.
- If the notebook is in edit mode, press `esc` on the keyboard to change to command mode.

It can be tricky at first understanding the difference between the two modes, but you will get the hang of it over time!

## Shortcuts

After a few lessons, you will mind might end up moving quicker than you can edit the notebook. Learning a few shortcuts can help you speed up your editing process. You can find most of the notebook book shortcuts by selecting `Help` in the menu bar and clicking `Notebook Shortcuts`.

A few of the most common shortcuts that we use are:

1. Hold `shift` while pressing `return` on the keyboard: This shortcut runs the selected cell and moves to the next cell.
2. While in command mode, pressing `b` on the keyboard will create a new cell below the selected cell.
3. While in command mode, pressing `d` on the keyboard twice will delete the selected cell.
4. While in command mode, pressing `m` on the keyboard turns the selected cell into a Markdown cell.

## The Kernel

The kernel is a program that executes the code inside your notebook and outputs the results. We running a Python 3 kernel in this course. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

As we described above, next to every code cell, you'll see some text that reads `[ ]:`. Before you run the cell, you'll see `[ ]:`. When the cell is running, you'll see `In [*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `In [1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

### Restarting the Kernel

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:

* At the top of your screen, click `Kernel` in the menu bar, then `Interrupt Kernel`.
* If that doesn't help, click `Kernel` in the menu bar, then `Restart Kernel`. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.
* If that doesn't help, restart your server by completing the following tasks in order:
    1. Save your work by clicking `File` in the menu bar, then `Save Notebook`. 
    2. Click `File` again, then`Hub Control Panel`. 
    3. Choose `Stop My Server` to shut it down, then `Start My Server` to start it back up.
    4. Open the notebook using the link from the Canvas assignment. You'll still have to run your code cells again.
    
Note that if your notebook is inactive for a while, then the kernel will shutdown on its own and you will have to run your code cells from the start of your notebook up until where you paused your work.

### Running Many Cells Automatically

There will be times when you need to re-run your cells. For example, you stopped working on an assignment and came back to it later, but the kernel had already shut down. If you want to quickly catch up to where last were, then 

1. Select the cell you last worked on.
2. Select `Kernel` from the menu bar, then `Restart Kernel and Run up to Selected Cell ...`.

Note that the cells will run in order from top to bottom and stop by executing the selected cell **unless** the kernel encounters an error. If there is an error in any of the code cells above the one you selected, then the process will stop at that point. You should probably address that error before moving one since most of our work depends on what has been done in previous cells.

## Errors

Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:

1. The rules are relatively _simple_.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are _rigid_.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

**Errors are okay; even experienced programmers make many errors.** When you make an error, you just have to find the source of the problem, fix it, and move on.

We have made an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

Depending on the version of Python being used to run the notebook, the error might look slightly different. We are using version 3.10 of Python, so the error should show `SyntaxError: incomplete input` at the bottom of the output. Additionally, there should be a ^ symbol attempting to point to the issue. In this case, the issue is that there is a missing parenthesis.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it. (Of course, if you're frustrated, ask for help.)

### Task 4 📍🔎

<!-- BEGIN QUESTION -->

For this task, complete the following:

1. Edit the above cell containing the code `print("This line is missing something."` to fix the `SyntaxError`.
2. Re-run the cell and notice that the error message is replaced with the correct output.
3. In the following Markdown cell, explain what you did to fix the error.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Check out [the official Python 3.10 documentation on Exceptions](https://docs.python.org/3.10/library/exceptions.html#concrete-exceptions) to view the types of errors you could see. Note that this documentation is a bit overwhelming to newcomers

## Submitting Your Lab to Canvas

Once you have finished working on the lab questions, prepare to submit your work in Canvas by completing the following steps:

1. Double-check that you have responded to all the open-response questions.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `grader.check_all()`. This command will run all of the run tests on all your responses to the code questions.
3. Select the menu item `File` and `Save and Checkpoint` in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items `File`, `Save and Export Notebook As ...`, and `HTML` in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment, click Start Assignment or New Attempt to upload the downloaded HTML file.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_2 (points: 1)
```
>>> 
>>> 'grader' in vars()
True
```
failure_message: ❌ It doesn't seem that you successfully ran the cell starting with # Initialize Otter.
success_message: ✅ You've successfully run the cell and initialized Otter!
