```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

<div style="text-align: center;">
    <h1>Homework 7: Testing Hypotheses</h1>
    <em>View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.</em>
</div>

**Reading**: 
* [Testing Hypotheses](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)
* [Comparing Two Samples](https://inferentialthinking.com/chapters/12/Comparing_Two_Samples.html)

Please complete this notebook by filling in the cells provided.

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Spam Calls: 781


Yanay gets a lot of spam calls. An area code is defined to be a three digit number from 200-999 inclusive. In reality, many of these area codes are not in use, but for this question we'll simplify things and assume they all are. **Throughout these questions, you should assume that Yanay's area code is 781.**

### Task 01 📍

Assuming each area code is just as likely as any other, what's the probability that the area code of two back to back spam calls are 781?



_Points:_ 2

```python
prob_781 = ...
prob_781
```

```python
grader.check("task_01")
```

### Task 02 📍

Rohan already knows that Yanay's area code is 781. Rohan randomly guesses the last 7 digits (0-9 inclusive) of his phone number. What's the probability that Rohan correctly guesses Yanay's number, assuming he’s equally likely to choose any digit?

*Note: A phone number contains an area code and 7 additional digits, i.e. xxx-xxx-xxxx*


_Points:_ 2

```python
prob_yanay_num = ...
prob_yanay_num
```

```python
grader.check("task_02")
```

Yanay suspects that there's a higher chance that the spammers are using his area code (781) to trick him into thinking it's someone from his area calling him. Ashley thinks that this is not the case, and that spammers are just choosing area codes of the spam calls at random from all possible area codes (*Remember, for this question we’re assuming the possible area codes are 200-999, inclusive*). Yanay wants to test his claim using the 50 spam calls he received in the past month.

Here's a dataset of the area codes of the 50 spam calls he received in the past month. Run the following code cell to load the data.

```python
spam = Table().read_table('spam.csv')
spam
```

### Task 03 📍🔎

<!-- BEGIN QUESTION -->

Define the null hypothesis and alternative hypothesis for this investigation. 

*Remember that your null hypothesis should fully describe a probability model that we can use for simulation later.*



_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 04 📍

Which of the following test statistics would be a reasonable choice to help differentiate between the two hypotheses?

For a refresher on choosing test statistics, check out the textbook section on [Test Statistics](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-2-the-test-statistic).

1. The proportion of area codes that are 781 in 50 random spam calls
2. The total variation distance (TVD) between probability distribution of randomly chosen area codes, and the observed distribution of area codes. (*Remember the possible area codes are 200-999 inclusive*)
3. The probability of getting an area code of 781 out of all the possible area codes.
4. The proportion of area codes that are 781 in 50 random spam calls divided by 2
5. The number of times you see the area code 781 in 50 random spam calls

Assign `reasonable_test_statistics` to an array of numbers corresponding to these test statistics.


_Points:_ 3

```python
reasonable_test_statistics = ...
```

```python
grader.check("task_04")
```

---

### 💡 Test Statistic

_For the rest of this section, use the number of times you see the area code 781 in 50 spam calls as your test statistic._

---

### Task 05 📍

Write a function called `simulate` that generates exactly one simulated value of your test statistic under the null hypothesis.  It should take no arguments and simulate 50 area codes under the assumption that the result of each area is sampled from the range 200-999 inclusive with equal probability. Your function should return the number of times you saw the 781 area code in those 50 random spam calls.



_Points:_ 3

```python
possible_area_codes = ...

def simulate():
    ...
    
# Call your function to make sure it works
simulate()
```

```python
grader.check("task_05")
```

### Task 06 📍

Using `simulate`, generate 20,000 simulated values of the number of times you see the area code 781 in 50 random spam calls. Assign `test_statistics_under_null` to an array that stores the result of each of these trials. 

_Points:_ 1

```python
repetitions = ...
test_statistics_under_null = ...

...
    
test_statistics_under_null
```

```python
grader.check("task_06")
```

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Using the values in `test_statistics_under_null`, generate a histogram of the empirical distribution of the number of times you saw the area code 781 in your simulation. 

_Use the provided bins when making the histogram._


_Points:_ 2

```python
bins = np.arange(0,5,1) # Use these provided bins

...
```

<!-- END QUESTION -->

### Task 08 📍

1. Calculate the observed value of the test statistic from the `spam` table and name the resulting value `observed_val`.
2. Using the simulated statistics, compute an empirical P-value for this test and name the result `p_value_1`.



_Points:_ 2

```python
observed_val = ...
p_value_1 = ...
p_value_1
```

```python
grader.check("task_08")
```

### Task 09 📍🔎

<!-- BEGIN QUESTION -->

Suppose you use a P-value cutoff of 1%. 

What do you conclude from the hypothesis test? 

Make sure your response addresses a decision between the null and alternative hypotheses, and briefly explain you reasoning.


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Spam Calls: Multiple Spammers

Instead of checking if the area code is equal to his own, Yanay decides to check if the area code matches the area code of one of the 8 places he's been to recently, and wants to test if it's more likely to receive a spam call with an area code from any of those 8 places. These are the area codes of the places he's been to recently: 781, 617, 509, 510, 415, 858, 339, 626.

### Task 10 📍

<!-- BEGIN QUESTION -->

Define the null hypothesis and alternative hypothesis for this investigation.

*Remember that your null hypothesis should fully describe a probability model that we can use for simulation later.*



_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### 💡 Test Statistic

_For the rest of this section, use the number of times you see any of the area codes of the places Yanay has been to in 50 spam calls as your test statistic._

---

### Task 11 📍

Write a function called `simulate_visited_area_codes` that generates exactly one simulated value of your test statistic under the null hypothesis.  It should take no arguments and simulate 50 area codes under the assumption that the result of each area is sampled from the range 200-999 inclusive with equal probability. Your function should return the number of times you saw any of the area codes of the places Yanay has been to in those 50 spam calls.

_You may find the textbook [section](https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html) on the `sample_proportions` function to be useful._


_Points:_ 3

```python
model_proportions = ...
def simulate_visited_area_codes():
    ...
    
# Call your function to make sure it works
simulate_visited_area_codes()
```

```python
grader.check("task_11")
```

### Task 12 📍

Using `simulate_visited_area_codes`, generate 20,000 simulated values of the number of times you see any of the area codes of the places Yanay has been to in 50 random spam calls. Store the result of each of these trials in the array `test_statistics_under_null`. 

_Points:_ 1

```python
visited_test_statistics_under_null = ...

repetitions = ...
...

visited_test_statistics_under_null
```

```python
grader.check("task_12")
```

### Task 13 📍🔎

<!-- BEGIN QUESTION -->

Using the results stored in `visited_test_statistics_under_null`, generate a histogram of the empirical distribution of the number of times you saw any of the area codes of the places Yanay has been to in your simulation. 

_Use the provided bins when making the histogram._


_Points:_ 2

```python
bins_visited = np.arange(0,6,1) # Use these provided bins
...
```

<!-- END QUESTION -->

### Task 14 📍

1. Set `visited_observed_value` to be the observed value of the test statistic.
2. Using the values in `visited_test_statistics_under_null`, compute an empirical P-value for this test and name the results `p_value_2`.



_Points:_ 2

```python
visited_area_codes = make_array(781, 617, 509, 510, 415, 858, 339, 626)
visited_observed_value = ...
p_value_2 = ...
p_value_2
```

```python
grader.check("task_14")
```

### Task 15 📍🔎

<!-- BEGIN QUESTION -->

Suppose you use a P-value cutoff of 0.05% (**Note: that’s 0.05%, not our usual cutoff of 5%**). What do you conclude from the hypothesis test? Why?


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 16 📍

Set `section_2_reflection_1` to be the integer corresponding to the best description of what `p_value_2` represents in the context of this problem.

1. The probability that the spam calls favored the visited area codes.
2. The probability that they didn't favor.
3. Neither 1 or 2.

_Points:_ 1

```python
section_2_reflection_1 = ...
```

```python
grader.check("task_16")
```

### Task 17 📍

<!-- BEGIN QUESTION -->

Set `section_2_reflection_2` to be the integer corresponding to the best description of what 0.05% (the P-value cutoff) represents in the context of this problem.

1. The probability that the spam calls favored the visited area codes.
2. The probability that they didn't favor. 
3. The probability of the test rejecting the null hypothesis if the null hypothesis were true.
4. Neither 1, 2, or 3.

_Points:_ 1

```python
section_2_reflection_2 = ...
```

```python
grader.check("task_17")
```

<!-- END QUESTION -->

### Task 18 📍🔎

<!-- BEGIN QUESTION -->

Suppose you run this test for 4000 different people after observing each person's last 50 spam calls. When you reject the null hypothesis for a person, you accuse the spam callers of favoring the area codes that person has visited. If the spam callers were not actually favoring area codes that people have visited, can we compute how many times we will incorrectly accuse the spam callers of favoring area codes that people have visited? If so, what is the number? Explain your answer. Assume a 0.05% P-value cutoff.


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Spam Calls: Practice with A/B Tests

Yanay collects information about this month's spam calls. The table `with_labels` is a sampled table, where the `Area Code Visited` column contains either `"Yes"` or `"No"` which represents whether or not Yanay has visited the location of the area code. The `Picked Up` column is `1` if Yanay picked up and `0` if he did not pick up.

Run the following cell to load the data.

```python
with_labels = Table().read_table("spam_picked_up.csv")
with_labels
```

Yanay is going to perform an A/B Test to see whether or not he is more likely to pick up a call from an area code he has visited. Specifically, his null hypothesis is that there is no difference in the distribution of calls he picked up between visited and not visited area codes, with any difference due to chance. His alternative hypothesis is that there is a difference between the two categories, specifically that he thinks that he is more likely to pick up if he has visited the area code. We are going to perform a [permutation test](https://www.inferentialthinking.com/chapters/12/1/AB_Testing.html#Permutation-Test) to test this. Our test statistic will be the difference in proportion of calls picked up between the area codes Yanay visited and the area codes he did not visit.

### Task 19 📍

1. Complete the `difference_in_proportion` function to have it calculate this test statistic.
2. Use the function to find the observed value. 

_The function takes in a sampled table which can be any table that has the same columns as `with_labels`. We'll call `difference_in_proportion` with the sampled table `with_labels` in order to find the observed difference in proportion._


_Points:_ 2

```python
def difference_in_proportion(sample):
    # Take a look at the code for `proportion_visited` and use that as a 
    # hint of what `proportions` should be assigned to
    proportions = ...
    proportion_visited = proportions.where("Area Code Visited", "Yes").column("Picked Up mean").item(0)
    proportion_not_visited = proportions.where("Area Code Visited", "No").column("Picked Up mean").item(0)
    ...

observed_diff_proportion = difference_in_proportion(with_labels)
observed_diff_proportion
```

```python
grader.check("task_19")
```

### Task 20 📍

To perform a permutation test we shuffle the labels, because our null hypothesis is that the labels don't matter because the distribution of calls he picked up between visited and not visited area codes come from same underlying distribution. The labels in this case is the `"Area Code Visited"` column containing `"Yes"` and `"No"`.

Using `difference_in_proportion`, create a function `simulate_one_stat` to shuffle the table and return a test statistic.

_Remember that to shuffle labels, we sample without replacement and then replace the appropriate column with the new shuffled column._


_Points:_ 1

```python
def simulate_one_stat():
    shuffled = ...
    original_with_shuffled_labels = ...
    return difference_in_proportion(original_with_shuffled_labels)

one_simulated_test_stat = simulate_one_stat() 
one_simulated_test_stat
```

```python
grader.check("task_20")
```

### Task 21 📍🔎

<!-- BEGIN QUESTION -->

Using `simulate_one_stat`, generate 1,000 simulated test statistic values. Store the result of each of these trial in the array `test_stats`.

We also provided code that'll generate a histogram for you after generating a 1000 simulated test statistic values.

_Points:_ 2

```python
trials = ...
test_stats = ...

...

# This generates a histogram of values and the red dot is the observed value
Table().with_column("Simulated Proportion Difference", test_stats).hist("Simulated Proportion Difference");
plt.plot(observed_diff_proportion, 0, 'ro', markersize=15);
```

<!-- END QUESTION -->

### Task 22 📍

Using `test_stats`, compute the empirical p-value for this test, and assign that value to `p_value_3`.


_Points:_ 2

```python
p_value_3 = ...
p_value_3
```

```python
grader.check("task_22")
```

For `p_value_ab`, you should be getting a value around 10-15%. If our p-value cutoff is 5%, the data is more consistent with the null hypothesis - that there is no difference in the distribution of calls Yanay picked up between visited and not visited area codes.

# Submit your Homework to Canvas

Once you have finished working on the homework questions, prepare to submit your work in Canvas by completing the following steps.

1. Double-check that you have run the code cell near the end of the notebook that contains the command "grader.check_all()". This command will run all of the run tests on all your responses to the code questions.
2. Double-check that you have responded to all the open-response questions.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download Notebook" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> 
>>> 0 <= prob_781 <= 1
True
```
failure_message: ❌ prob_781 should be a number between 0 and 1.
success_message: ✅ prob_781 is a number between 0 and 1.


Tests for task_02 (points: 2)
```
>>> 
>>> 0 <= prob_yanay_num <= 1
True
```
failure_message: ❌ prob_yanay_num should be a number between 0 and 1.
success_message: ✅ prob_yanay_num is a number between 0 and 1.


Tests for task_04 (points: 3)
```
>>> 
>>> type(reasonable_test_statistics) == np.ndarray
True
```
failure_message: ❌ reasonable_test_statistics should be an array.
success_message: ✅ reasonable_test_statistics is an array.
```
>>> 
>>> 5 in set(reasonable_test_statistics) and len(reasonable_test_statistics) > 0
True
```
failure_message: ❌ reasonable_test_statistics should at least contian 5.
success_message: ✅ reasonable_test_statistics at least contains 5.


Tests for task_05 (points: 3)
```
>>> 
>>> np.std([simulate() for _ in range(1000)]) > 0
True
```
failure_message: ❌ It looks like your simulation isn't random.
success_message: ✅ Your simulation appears to be producing random values.


Tests for task_06 (points: 1)
```
>>> 
>>> len(test_statistics_under_null) == 20000
True
```
failure_message: ❌ test_statistics_under_null should contain 20,000 values.
success_message: ✅ test_statistics_under_null contains 20,000 values.


Tests for task_08 (points: 2)
```
>>> 
>>> 0 <= p_value_1 < 1
True
```
failure_message: ❌ p_value should be a value between 0 and 1.
success_message: ✅ p_value is a value between 0 and 1.


Tests for task_11 (points: 3)
```
>>> 
>>> np.std([simulate_visited_area_codes() for _ in range(1000)]) > 0
True
```
failure_message: ❌ Your simulation should be producing random values.
success_message: ✅ Your simulation appears to be producing random values.
```
>>> 
>>> model_proportions.item(0) + model_proportions.item(1) == 1
True
```
failure_message: ❌ The sum of the items in model proportions should be 1.
success_message: ✅ The sum of the items in model proportions is 1.


Tests for task_12 (points: 1)
```
>>> 
>>> len(visited_test_statistics_under_null) == 20000
True
```
failure_message: ❌ visited_test_statistics_under_null should contain 20,000 values.
success_message: ✅ visited_test_statistics_under_null contains 20,000 values.


Tests for task_14 (points: 2)
```
>>> 
>>> 0 <= p_value_2 < 1
True
```
failure_message: ❌ p_value_2 should be a number between 0 and 1.
success_message: ✅ p_value_2 is a number between 0 and 1.


Tests for task_16 (points: 1)
```
>>> 
>>> section_2_reflection_1 in {1, 2, 3}
True
```
failure_message: ❌ section_2_reflection_1 should name 1, 2, or 3.
success_message: ✅ section_2_reflection_1 is a possible value.


Tests for task_17 (points: 1)
```
>>> 
>>> section_2_reflection_2 in {1, 2, 3, 4}
True
```
failure_message: ❌ section_2_reflection_2 should name 1, 2, 3, or 4.
success_message: ✅ section_2_reflection_2 is a possible value.


Tests for task_19 (points: 2)
```
>>> 
>>> -1 <= observed_diff_proportion <= 1
True
```
failure_message: ❌ observed_diff_proportion should be a number between -1 and 1, inclusive.
success_message: ✅ observed_diff_proportion is a number between -1 and 1, inclusive.
```
>>> 
>>> np.round(observed_diff_proportion, 3) == 0.219
True
```
failure_message: ❌ The observed difference in proportion should be about 0.219
success_message: ✅ The observed difference in proportion is approximately 0.219


Tests for task_20 (points: 1)
```
>>> 
>>> -0.75 <= one_simulated_test_stat <= 0.75
True
```
failure_message: ❌ one_simulated_test_stat should be a value between -0.75 and 0.75, inclusive.
success_message: ✅ one_simulated_test_stat is a value between -0.75 and 0.75, inclusive.


Tests for task_22 (points: 2)
```
>>> 
>>> 0 <= p_value_3 < 1
True
```
failure_message: ❌ p_value_3 should be a number between 0 and 1.
success_message: ✅ p_value_3 is a number between 0 and 1.
